<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="coderunner">
    <name>
      <text>Teste jsdom (Node.js)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Escreve a função <code>titulo(html)</code> que usa o <b>jsdom</b> para devolver o texto do primeiro <code>&lt;h1&gt;</code> do HTML recebido.</p>
<p>Exemplo: <code>titulo('&lt;h1&gt;Olá&lt;/h1&gt;')</code> devolve <code>'Olá'</code>.</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>nodejs</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>1</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>12</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[const { JSDOM } = require('jsdom');

function titulo(html) {
  // ...
}
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[const { JSDOM } = require('jsdom');

function titulo(html) {
  const dom = new JSDOM(html);
  const h1 = dom.window.document.querySelector('h1');
  return h1 ? h1.textContent : null;
}
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs>10</cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode>
          <text>console.log(titulo('&lt;h1&gt;Olá&lt;/h1&gt;'));</text>
        </testcode>
        <stdin>
          <text></text>
        </stdin>
        <expected>
          <text>Olá</text>
        </expected>
        <extra>
          <text></text>
        </extra>
        <display>
          <text>SHOW</text>
        </display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode>
          <text>console.log(titulo('&lt;div&gt;&lt;p&gt;x&lt;/p&gt;&lt;h1&gt;DEISI&lt;/h1&gt;&lt;h1&gt;outro&lt;/h1&gt;&lt;/div&gt;'));</text>
        </testcode>
        <stdin>
          <text></text>
        </stdin>
        <expected>
          <text>DEISI</text>
        </expected>
        <extra>
          <text></text>
        </extra>
        <display>
          <text>SHOW</text>
        </display>
      </testcase>
    </testcases>
  </question>
</quiz>
